## Exercise 04.1 (simple function)

Write a function called `is_even` which takes an integer as an argument and returns `True` if the argument is even, and otherwise returns `False`. Test your function for several values.

<!-- grade_id: cell-f07e8ac1f9838913, points: 0, grade: false, solution: true -->
```python
def is_even(x):
    ...
    if x % 2 == 0:
        return True # all even numebrs are divisible by 2
    else:
        return False
    
```

<!-- grade_id: cell-8a4ca795d9fa96dd, points: 0, grade: true, solution: false -->
```python
## tests ##
assert is_even(0) == True
assert is_even(101) == False
assert is_even(982) == True 
assert is_even(-5) == False
assert is_even(-8) == True
```

## Exercise 04.2 (functions and default arguments)

Write a single function named `magnitude` that takes each component of a vector of length 2 or 3 and returns the magnitude.
Use default arguments to handle vectors of length 2 or 3 with the same code. Test your function for correctness against hand calculations for a selection of values.

<!-- grade_id: cell-693a09b45c6bb532, points: 0, grade: false, solution: true -->
```python
import math
...
def magnitude(x,y,z=0): # Unless othewise inputted, a zero value is given for x, y or z
    return ((x**2)+(y**2)+(z**2))**0.5
print(magnitude(2,3)) # Testing the concept
```

Output:
```
3.605551275463989
```

<!-- grade_id: cell-384a85ba97b226d7, points: 0, grade: true, solution: false -->
```python
## tests ##
assert math.isclose(magnitude(3, 4), 5.0)
assert math.isclose(magnitude(4, 3), 5.0)
assert math.isclose(magnitude(4, 3, 0.0), 5.0)
assert math.isclose(magnitude(4, 0.0, 3.0), 5.0)
assert math.isclose(magnitude(3, 4, 4), 6.403124237)
```

## Exercise 04.3 (functions)

Given the coordinates of the vertices of a triangle, $(x_0, y_0)$, $(x_1, y_1)$ and $(x_2, y_2)$, the area $A$ of the triangle is given by:
$$
A = \left| \frac{x_0(y_1  - y_2) + x_1(y_2 - y_0) + x_2(y_0 - y_1)}{2} \right|
$$
Write a function named `area` that computes the area of a triangle given the coordinates of the vertices.
The order of the function arguments must be (`x0, y0, x1, y1, x2, y2)`.

Test the output of your function against some known solutions.

<!-- grade_id: cell-ed02439d4509b2c1, points: 0, grade: false, solution: true -->
```python
...
def area(x0,y0,x1,y1,x2,y2):
    return abs(((x0*(y1 - y2))+(x1*(y2-y0))+(x2*(y0 - y1))/2))  # equation of A
```

<!-- grade_id: cell-64f8b6708595aed3, points: 0, grade: true, solution: false -->
```python
## tests ##
x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 0.0
A = area(x0, y0, x1, y1, x2, y2)
assert math.isclose(A, 3.0)
```

## Exercise 04.4 (recursion)

The factorial of a non-negative integer $n$ is expressed recursively by:
$$
n! = 
\begin{cases}
1 & n = 0 \\
(n - 1)! \,n & n > 0
\end{cases}
$$

Develop a function named `factorial` for computing the factorial using recursion.
Test your function against the `math.factorial` function, e.g.

```python
import math
print("Reference factorial:", math.factorial(5))
```

Output:
```
Reference factorial: 120
```

<!-- grade_id: cell-9591c1fc1c49c167, points: 0, grade: false, solution: true -->
```python
...
def factorial(n): # integer n being the factorial in Qn (n!)
    if n == 0:
        return 1 # deals with 0, otherwise there would be an error
    elif n == 1:
        return 1 # deals with 1, otherwisw there would be an error
    else:
        return ((n)*factorial(n-1)) # multiplies the factorial of the previous number by the current number, giving the factorial of thus
print("Factorial of 5:", factorial(5))

import math
print("Reference value of factorial of 5:", math.factorial(5))
```

Output:
```
Factorial of 5: 120
Reference value of factorial of 5: 120
```

<!-- grade_id: cell-ffe589e9fb2f1419, points: 0, grade: true, solution: false -->
```python
## tests ##
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(2) == 2
assert factorial(5) == 120

import math
assert factorial(32) == math.factorial(32)
```

## Exercise 04.5 (functions and passing functions as arguments)

Restructure your program from the bisection problem in Exercise 02 to 

- Use a Python function to evaluate the mathematical function $f$ that we want to find the root of; 

and then

- Encapsulate the bisection algorithm inside a Python function, which takes as arguments:
  1. the function we want to find the roots of
  1. the points $x_{0}$ and $x_{1}$ between which we want to search for a root
  1. the tolerance for exiting the bisection algorithm (exit when $|f(x)| < \text{tol}$)
  1. maximum number of iterations (the algorithm should exit once this limit is reached)

For the first step, create a Python function for evaluating $f$, e.g.:
```python
def f(x):
    # Put body of the function f(x) here, returning the function value
```           
For the second step, encapsulate the bisection algorithm in a function:
```python
def compute_root(f, x0, x1, tol, max_it):

    # Implement bisection algorithm here, and return when tolerance is satisfied or
    # number of iterations exceeds max_it

    # Return the approximate root, value of f(x) and the number of iterations
    return x, f, num_it

# Compute approximate root of the function f
x, f_x, num_it = compute_root(f, x0=0, x1=1, tol=1.0e-6, max_it=1000)
```

You can try testing your program for a function $f(x)$ that is simpler from the function Exercise 02, e.g. $f(x) = x - 4$. 
A quadratic function, the roots of which you can find analytically, would be a good test case.

### Solution

Define the function for computing $f(x)$:

```python
import math

def my_f(x):
    """Evaluate polynomial function"""
    return ((x**5) / (10)) + (x**3) - (10 * (x**2)) + (4 *( x) + 7)
```

<!-- grade_id: cell-622edcb129dcc7f2, points: 0, grade: false, solution: true -->
```python
# Create the function that performs the bisection:
def compute_root(f, x0, x1, tol, max_it):
    """Compute roots of a function using bisection"""
    
    error = tol + 1.0

    it = 0
    while error > tol: # Iterates until tolerance is met
        x_mid = (x0 + x1) / 2

        f0 = my_f(x0)
        f_mid = my_f(x_mid)
    
        if f0*f_mid<0:
            x1 = x_mid # change in sign, thus value of x decreases to find the point of sign-change
        else:
            x0 = x_mid # x0 increases since there has been no change in sign, and f(x) is positive when x=0
        error = abs(f_mid)
        
        
        if it > max_it:
            print("Iteration count too high. Breaking out of loop")
            break
        print(it, x_mid, error)
        it += 1
    return x_mid, f_mid, it


```

<!-- grade_id: cell-cb5a5cdf1d790edb, points: 0, grade: true, solution: false -->
```python
## tests ##

x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=1000)

# Test solution for function in Exercise 02
assert math.isclose(x, 1.1568354368209839)
```

Output:
```
0 1.0 2.0999999999999996
1 1.5 5.365625000000001
2 1.25 1.36669921875
3 1.125 0.4477813720703132
4 1.1875 0.4408627510070797
5 1.15625 0.008327355980872753
6 1.171875 0.21507753478363156
7 1.1640625 0.1030741602036862
8 1.16015625 0.04729774910292761
9 1.158203125 0.019466230897839054
10 1.1572265625 0.005564689502373099
11 1.15673828125 0.0013825210451265946
12 1.156982421875 0.0020907873792292975
13 1.1568603515625 0.00035405894194262544
14 1.15679931640625 0.0005142496094645566
15 1.156829833984375 8.009997302949046e-05
16 1.1568450927734375 0.00013697832466341708
17 1.1568374633789062 2.8438885864900953e-05
18 1.1568336486816406 2.5830616070976475e-05
19 1.1568355560302734 1.304116775457942e-06
20 1.156834602355957 1.2263254177469207e-05
21 1.1568350791931152 5.479569834321296e-06
22 1.1568353176116943 2.087726812760593e-06
23 1.1568354368209839 3.9180508970559913e-07
```

#### Optional extension

Use recursion to write a `compute_root` function that *does not* require a `for` or `while` loop.

<!-- grade_id: cell-c944bb2229317d4d, points: 0, grade: false, solution: true -->
```python
def compute_root(f, x0, x1, tol, max_it, it_count=0):
    ...
    tol = 1.0e-6
error = tol + 1.0
x0 = 0
x1 = 2

counter = 0
    
    x_mid = (x0 + x1) / 2

    f0 = my_f(x0)
   
    f_mid = my_f(x_mid)
    if (f0)*(f_mid)<0:
        x1=x_mid # change in sign, thus value of x decreases to find the point of sign-change
    else:
        x0=x_mid # x0 increaseZs since there has been no change in sign, and f(x) is positive when x=0
    counter += 1
    # Call compute_root recursively
    return compute_root(f, x0, x1, tol, max_it, it_count)

```

Output:
```
Cell In[13], line 10
    x_mid = (x0 + x1) / 2
    ^
IndentationError: unexpected indent
```

<!-- grade_id: cell-9ebe2f13b2899a8f, points: 0, grade: true, solution: false -->
```python
## tests ##
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=1000)
assert math.isclose(x, 1.1568354368209839)
```

```python

```