## Exercise 05.1 (random numbers)

- Using the `randint` function from the `random` module (https://docs.python.org/3/library/random.html#random.randint) to
  develop a function `dice_roll` that emulates the roll of a dice with $n$ sides. The number of sides `n` should an argument to the function.

- For $n=6$, devise and implement a test to check that it is a fair dice.

#### (a) Dice roll code:

<!-- grade_id: cell-70a83466894475ee, points: 0, grade: false, solution: true -->
```python
import random
from random import randint
...
def dice_roll(n): # n is the input for the range
    for i in range(n):
        return randint(1,n) # comes up with any number within the range
print(dice_roll(100)) # visualising the test for personal reference
```

Output:
```
94
```

<!-- grade_id: cell-0970743b0c035f78, points: 0, grade: true, solution: false -->
```python
## tests ##
for n in range(1, 20):
    for j in range(100):
        value = dice_roll(n) 
        assert value >= 1 and value <= n
```

#### (b) Test for fairness

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
# To check for fairness, we roll the dice a large number of times and check how many times we get each value. 
# We expect to get each value close to $1/6$ of the times. 

...

def fairness_test(n):
    value_list = []
    for k in range(n):
        value = dice_roll(6)
        value_list.append(value)
  #  print(value_list)
    # I realise now that I could format this more compactly with an iteration
    print("number of 1s: "+ str(value_list.count(1))+", and its proportion in the test is: " + str(value_list.count(1)/n)) 
    print("number of 2s: "+ str(value_list.count(2))+", and its proportion in the test is: " + str(value_list.count(2)/n))
    print("number of 3s: "+ str(value_list.count(3))+", and its proportion in the test is: " + str(value_list.count(3)/n))
    print("number of 4s: "+ str(value_list.count(4))+", and its proportion in the test is: " + str(value_list.count(4)/n))
    print("number of 5s: "+ str(value_list.count(5))+", and its proportion in the test is: " + str(value_list.count(5)/n))
    print("number of 6s: "+ str(value_list.count(6))+", and its proportion in the test is: " + str(value_list.count(6)/n))
    for ik in range(6):
        if 0.15 < value_list.count(ik)/n < 0.2167:
            print("Proportionate result for " + str(ik+1))
        else:
            print("disproportionate result for " + str(ik+1))

fairness_test(1000)
```

Output:
```
number of 1s: 161, and its proportion in the test is: 0.161
number of 2s: 150, and its proportion in the test is: 0.15
number of 3s: 172, and its proportion in the test is: 0.172
number of 4s: 176, and its proportion in the test is: 0.176
number of 5s: 176, and its proportion in the test is: 0.176
number of 6s: 165, and its proportion in the test is: 0.165
disproportionate result for 1
Proportionate result for 2
disproportionate result for 3
Proportionate result for 4
Proportionate result for 5
Proportionate result for 6
```

## Exercise 05.2 (variance estimation)

For a random variable $X$, the variance of $X$ is defined as  

$$
\begin{align}
\mathrm{Var}\left[ X \right] &= \mathrm{E}\left[\left( X - \mu \right)^2\right]  \\ 
&= \mathrm{E}\left[ X^{2} \right] - \mathrm{E}\left[ X \right]^2
\end{align}
$$

where $\mathrm{E}$ is the 'expectation' (mean of something) and $\mu = \mathrm{E}(X)$ is the mean of $X$. If we have all data (the entire 'population'),
the variance can be computed from:

$$
\mathrm{Var}\left[ X \right] = \frac{\sum_{i=0}^{N-1} x^{2}_{i}}{N} - \left( \frac{ \left( \sum_{i=0}^{N-1} x_{i} \right)}{N} \right)^{2}
$$

Often, we only have a sample of data. For example, we might want to estimate the variance in height for students at a university using just a random sample of students. 
When using a sample from a larger data set to estimate the variance, the above formula has a *bias*. Therefore, it is common to use the *unbiased* estimator

$$
s^{2} = \left( \frac{\sum_{i=0}^{n-1} x^{2}_{i}}{n} - \left( \frac{ \left( \sum_{i=0}^{n-1} x_{i} \right)}{n} \right)^{2} \right) \frac{n}{n-1}
$$

to estimate the variance. In this exercise we will use the unbiased estimator.

1. Create a function that returns the estimated variance for a list of numbers based on the above equation. Test your function using 1 million samples drawn from a Gaussian distribution with a mean of 10 and a standard deviation of 3. Use your crsid to seed the random number generator
   
   *See hint below on how to create the sample.*

2. For a sample drawn from a distribution with mean $5 \times 10^6$ and standard deviation $2.0$, estimate the variance using (i) your function for estimating the variance, and (ii) using the `variance` function from the Python `statistics` module. Comment on and explain any significant differences in the results from (i) and (ii).

### Hint: sampling from a distribution

The function `random.gauss` can be used to sample a Gaussian distribution with a specified mean and standard deviation (square root of the variance) *N* times, e.g.:

```python
import random

random.seed("gnw20")  # See the random number generator

mu = 10.0    # mean of the distribution
sigma = 3.0  # Standard deviation
x = [random.gauss(mu, sigma) for i in range(8)]
print(x)
```

Output:
```
[11.176693418313244, 12.08792069683015, 12.649052763696117, 9.845742673456446, 7.317196693264903, 6.876099163343957, 15.181269044169637, 12.844967463974331]
```

### Solution

1. Estimate variance

```python
import random
random.seed("gnw20")

...


def variance_function(mu,sigma,n):
    y = [random.gauss(mu, sigma) for i in range(n)]
    #print(y)

    Ex2 = 0
    Ex = 0 # will be squared later
    
    for j in range(0,n):
        Ex2 += ((y[j])**2)
        Ex += ((y[j]))
    variance = (n/(n-1))*(((Ex2)/n)-((Ex/n)**2))
    print("variance when n is " + str(n) + " is: " + str(variance))

variance_function(5.0e6,2,100)  
variance_function(5.0e6,2,300)
variance_function(5.0e6,2,1000)
variance_function(5.0e6,2,10000)
variance_function(5.0e6,2,100000)
    
    

```

Output:
```
variance when n is 100 is: 3.8786300505050506
variance when n is 300 is: 3.9820234113712374
variance when n is 1000 is: 4.555336586586587
variance when n is 10000 is: 3.789441444144414
variance when n is 100000 is: 4.128947539475394
```

2. Use the `statistics` module to estimate the variance, and compare the estimated variance using `statistics` to the estimated variance using your implementation for the variance estimation.

```python
import statistics
for i in range(1,7):
    a = [random.gauss(mu, sigma) for i in range(10**i)]
    var = statistics.variance(a)
    print("variance from statistics module when n is " + str(10**i) + " is: " + str(var)) # produces different results each time, but generally gets closer to the real value of 4 as the sample size increases
```

Output:
```
variance from statistics module when n is 10 is: 2.877560732426835
variance from statistics module when n is 100 is: 10.137042956289656
variance from statistics module when n is 1000 is: 8.99708150140705
variance from statistics module when n is 10000 is: 9.043732856078526
variance from statistics module when n is 100000 is: 9.00090186038833
variance from statistics module when n is 1000000 is: 9.00438995519197
```

## Exercise 05.03 (optional, parallel processing)

Almost all modern computer processing units have multiple processing *cores*. To utilise the full performance of a processor, operations need to be performed in *parallel*, i.e. each processing core is given a task to perform.

Parallel computing is a very rich and technical area. To help exploit multi-core systems there are libraries that support parallel processing. Below is a simple example using the Python `multiprocessing` library.

```python
import multiprocessing

# On some operaring systems, you may need to uncomment the below line
# multiprocessing.set_start_method('fork')

def f(task):
    """A function that print the input argument and the id for the process that executes the function"""
    print(f"Task index {task}, process id: {multiprocessing.current_process()}\n")
     
    return str(multiprocessing.current_process())
        
        
# Using 3 'processes', execute the function 'f' four times (each time with a different argument)  
with multiprocessing.Pool(processes=3) as p:
    procs = p.map(f, [0, 1, 2, 3])  # Call function mysort three times
    
    print("Returned data (a list)") 
    print(procs)
```

Output:
```
Task index 2, process id: <ForkProcess name='ForkPoolWorker-3' parent=999 started daemon>
Task index 0, process id: <ForkProcess name='ForkPoolWorker-1' parent=999 started daemon>
Task index 1, process id: <ForkProcess name='ForkPoolWorker-2' parent=999 started daemon>



Task index 3, process id: <ForkProcess name='ForkPoolWorker-3' parent=999 started daemon>

Returned data (a list)
["<ForkProcess name='ForkPoolWorker-1' parent=999 started daemon>", "<ForkProcess name='ForkPoolWorker-2' parent=999 started daemon>", "<ForkProcess name='ForkPoolWorker-3' parent=999 started daemon>", "<ForkProcess name='ForkPoolWorker-3' parent=999 started daemon>"]
```

Consider the below code that creates a list of lists of sorted numbers.

```python
def mysort(N):
    """Create a randomly ordered list of integers of length N, and return the sorted list"""
    # Create randomly ordered list
    x = random.sample(range(0, N), N)
 
    # Return sorted list of numbers
    return sorted(x)

def sorted_lists(N, p):
    """Create a list of sorted lists"""
    data = []
    for i in range(p):
        data.append(mysort(N))

%time x = sorted_lists(1000000, 5)
```

Output:
```
CPU times: user 6.91 s, sys: 434 ms, total: 7.34 s
Wall time: 7.72 s
```

Use the `multiprocessing` module to perform the above operation in parallel. Investigate how the processing time changes with the number of processes, and in particular the average time per 'create and sort operation' when changing the number of processes.

```python
...
```

Output:
```
Ellipsis
```