## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    ...
    if n == 0:
        return 0  # This doesn't call f, so it breaks out of the recursion loop, preventing an exception
    elif n == 1:
        return 1  # This doesn't call f, so it breaks out of the recursion loop, preventing an exception
    elif n < 0:
        raise ValueError("Negative number inputtted. Try again")
    else:
        return f(n - 1) + f(n - 2)  # This calls f for n-1 and n-2 (recursion), and returns the sum
         # This would not have worked with 0 and 1, since a negative number would have been called, leading to an exception
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##

# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7

def compute_root(f, x0, x1, tol, max_it):
    ...
    error = tol + 1.0

    it = 0
    while error > tol: # Iterates until tolerance is met
        x_mid = (x0 + x1) / 2

        f0 = my_f(x0)
        f_mid = my_f(x_mid)
    
        if f0*f_mid<0:
            x1 = x_mid # change in sign, thus value of x decreases to find the point of sign-change
        else:
            x0 = x_mid # x0 increases since there has been no change in sign, and f(x) is positive when x=0
        error = abs(f_mid)
        
        
        if it > max_it:
            raise RuntimeError("Too mamy iterations! Loop terminated") # gives an error if the number of iterations exceeds that of the maximum
        if tol == 0:
            raise RangeError("Don't be so intolerant! tolerance, tol must be a positive number") # makes sure tolerance will not lead to an infinite loop
        if tol < 0:
            raise RangeError("tolerance, tol must be a positive number")
        if it <= 0:
            raise RangeError("Iteration, it must be a positive integer")
        print(it, x_mid, error)
        it += 1
    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=20)
```

Output:
```
0 1.0 2.0999999999999996
1 1.5 5.365625000000001
2 1.25 1.36669921875
3 1.125 0.4477813720703132
4 1.1875 0.4408627510070797
5 1.15625 0.008327355980872753
6 1.171875 0.21507753478363156
7 1.1640625 0.1030741602036862
8 1.16015625 0.04729774910292761
9 1.158203125 0.019466230897839054
10 1.1572265625 0.005564689502373099
11 1.15673828125 0.0013825210451265946
12 1.156982421875 0.0020907873792292975
13 1.1568603515625 0.00035405894194262544
14 1.15679931640625 0.0005142496094645566
15 1.156829833984375 8.009997302949046e-05
16 1.1568450927734375 0.00013697832466341708
17 1.1568374633789062 2.8438885864900953e-05
18 1.1568336486816406 2.5830616070976475e-05
19 1.1568355560302734 1.304116775457942e-06
20 1.156834602355957 1.2263254177469207e-05
21 1.1568350791931152 5.479569834321296e-06
22 1.1568353176116943 2.087726812760593e-06
23 1.1568354368209839 3.9180508970559913e-07
0 1.0 2.0999999999999996
1 1.5 5.365625000000001
2 1.25 1.36669921875
3 1.125 0.4477813720703132
4 1.1875 0.4408627
```